
# Lab 8 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 8, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import plotly.express as px
import statsmodels.api as sm
import os
from sklearn.preprocessing import Binarizer, QuantileTransformer, FunctionTransformer

import warnings
warnings.filterwarnings('ignore')
```

```python
# By setting a seed, we guarantee that we will see the same results each time we run this cell.
np.random.seed(23)

# Generates a random scatter plot
x = np.arange(1, 101) + np.random.normal(0, 0.5, 100)
y = 2 * ((x + np.random.normal(0, 1, 100)) ** 2) + np.abs(x) * np.random.normal(0, 30, 100)
df_1 = pd.DataFrame().assign(x=x, y=y)

px.scatter(df_1, x='x', y='y', trendline="ols", trendline_color_override="red")
```

```python
df_1['root y'] = np.sqrt(df_1['y'])

px.scatter(df_1, x='x', y='root y', trendline="ols", trendline_color_override="red")
```

```python
# By setting a seed, we guarantee that we will see the same results each time we run this cell
np.random.seed(32)

# Generates a different random scatter plot
x = np.linspace(2, 5, 100)
y = 10 * (np.e ** x) + np.abs(x) * np.random.normal(0, 5, 100) + np.random.normal(0, 30, 100)
df_2 = pd.DataFrame().assign(x=x, y=y)

px.scatter(df_2, x='x', y='y', trendline="ols", trendline_color_override="red")
```

```python
df_2['root y'] = np.sqrt(df_2['y'])

px.scatter(df_2, x='x', y='root y', trendline="ols", trendline_color_override="red")
```

```python
# Feel free to use this function directly to help you answer Question 1.
def create_residual_plot(df, x, y):
    df = df.copy()
    from sklearn.linear_model import LinearRegression
    model = LinearRegression()
    model.fit(df[[x]], df[y])
    df[f'{y} residuals'] = df[y] - model.predict(df[[x]])
    return px.scatter(df, x=x, y=f'{y} residuals', trendline='ols', trendline_color_override='red')

create_residual_plot(df_2, 'x', 'root y')
```

```python
df_2['log y'] = np.log(df_2['y'])

px.scatter(df_2, x='x', y='log y', trendline="ols", trendline_color_override="red")
```

```python
create_residual_plot(df_2, 'x', 'log y')
```

```python
homeruns_fp = os.path.join('data', 'homeruns.csv')
homeruns = pd.read_csv(homeruns_fp)
```

```python
grader.check("q1")
```

```python
diamonds = pd.read_csv(os.path.join('data', 'diamonds.csv'))
diamonds.head()
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(os.path.join('data', 'diamonds.csv'))
out_q2 = create_ordinal(diamonds)
```

```python
grader.check("q2")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(os.path.join('data', 'diamonds.csv'))
out1_q3 = create_one_hot(diamonds)
out2_q3 = create_proportions(diamonds)
```

```python
grader.check("q3")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(os.path.join('data', 'diamonds.csv'))
out_q4 = create_quadratics(diamonds)
```

```python
grader.check("q4")
```

```python
from sklearn.linear_model import LinearRegression

# X = ...
# y = ...

# lr = LinearRegression()
# lr.fit(X, y)  # X is a DataFrame of training data; y is a Series of prices
# lr.score(X, y)  # R-squared
# lr.predict(X) # predicted prices
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
import numbers
out_q5 = comparing_performance()
```

```python
grader.check("q5")
```

```python
from sklearn.preprocessing import Binarizer, QuantileTransformer, FunctionTransformer
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(os.path.join('data', 'diamonds.csv'))
q6a_trans = TransformDiamonds(diamonds)
q6a_out = q6a_trans.transform_carat(diamonds)
```

```python
grader.check("q6.1")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q6b_trans = TransformDiamonds(diamonds)
q6b_out = q6b_trans.transform_to_quantile(diamonds)
q6b_trans_top_1000 = TransformDiamonds(diamonds[:1000])
q6b_out_top_1000 = q6b_trans_top_1000.transform_to_quantile(diamonds)
```

```python
grader.check("q6.2")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(os.path.join('data', 'diamonds.csv'))
q6c_trans = TransformDiamonds(diamonds)
q6c_out = q6c_trans.transform_to_depth_pct(diamonds)
```

```python
grader.check("q6.3")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> best_transformation() in [1, 2, 3, 4]
True
```
```
>>> best_transformation() == 1
True
```
failure_message: root transformation


Tests for q2 (points: None)
```
>>> set(out_q2.columns) == {'ordinal_cut', 'ordinal_clarity', 'ordinal_color'}
True
```
```
>>> np.unique(out_q2['ordinal_cut']).tolist() == [0, 1, 2, 3, 4]
True
```
```
>>> out_q2.loc[1, 'ordinal_cut'] == 3
True
```
```
>>> set(out_q2['ordinal_cut']) == {0, 1, 2, 3, 4}
True
```
```
>>> set(out_q2['ordinal_clarity']) == {0, 1, 2, 3, 4, 5, 6, 7}
True
```
```
>>> set(out_q2['ordinal_color']) == {0, 1, 2, 3, 4, 5, 6}
True
```
```
>>> cond1 = np.isclose(out_q2['ordinal_cut'].mean(), 2.904097144975899, atol=0.01)
>>> 
>>> cond2 = np.isclose((4 - out_q2['ordinal_cut']).mean(), 2.904097144975899, atol=0.01)
>>> 
>>> cond1 or cond2
True
```
```
>>> cond1 = np.isclose(out_q2['ordinal_clarity'].mean(), 3.0510196514645904, atol=0.01)
>>> 
>>> cond2 = np.isclose((7 - out_q2['ordinal_clarity']).mean(), 3.0510196514645904, atol=0.01)
>>> 
>>> cond1 or cond2
True
```
```
>>> cond1 = np.isclose(out_q2['ordinal_color'].mean(), 2.5941972562106046, atol=0.05)
>>> 
>>> cond2 = np.isclose((6 - out_q2['ordinal_color']).mean(), 2.5941972562106046, atol=0.05)
>>> 
>>> cond1 or cond2
True
```


Tests for q3 (points: None)
```
>>> out1_q3.shape == (53940, 20)
True
```
```
>>> out1_q3.columns.str.startswith('one_hot').all()
True
```
```
>>> out1_q3.isin([0,1]).all().all()
True
```
```
>>> out2_q3.shape[1] == 3
True
```
```
>>> out2_q3.columns.str.startswith('proportion_').all()
True
```
```
>>> ((out2_q3 >= 0) & (out2_q3 <= 1)).all().all()
True
```
```
>>> np.isclose(out1_q3['one_hot_clarity_VVS2'].mean(), 0.0939191694475343, atol=0.01)
True
```
```
>>> np.isclose(out1_q3['one_hot_cut_Fair'].mean(), 0.029847979236188357, atol=0.01)
True
```
```
>>> out2_q3.shape == (53940, 3)
True
```
```
>>> np.isclose(out2_q3['proportion_cut'], 0.399537, atol=0.01).sum() == 21551
True
```
```
>>> out2_q3['proportion_color'].nunique() == 7
True
```


Tests for q4 (points: None)
```
>>> out_q4.columns.str.contains(' * ').all()
True
```
```
>>> ('x * z' in out_q4.columns) or ('z * x' in out_q4.columns)
True
```
```
>>> out_q4.shape[1] == 15
True
```
```
>>> out_q4.shape == (53940, 15)
True
```
```
>>> cols = pd.read_csv(os.path.join('data', 'diamonds.csv')).columns
>>> s = set([f'{x} * {y}' for x in cols for y in cols])
>>> set(out_q4.columns) <= s
True
```
```
>>> np.isclose(out_q4.mean(axis=1).iloc[0], 310.24, atol=0.1)
True
```
```
>>> out_q4.max().max() == 5767
True
```
```
>>> out_q4.min().min() == 0
True
```
```
>>> out_q4.min().loc[lambda x: x > 0].min() == 11
True
```


Tests for q5 (points: None)
```
>>> len(out_q5) == 6
True
```
```
>>> isinstance(out_q5[0], numbers.Real)
True
```
```
>>> all(isinstance(x, str) for x in out_q5[2:-1])
True
```
```
>>> out_q5[1] > out_q5[-1]
True
```
```
>>> np.isclose(out_q5[0], 0.8493305264354857, atol=0.01)
True
```
```
>>> np.isclose(out_q5[1], 1548.53, atol=0.1)
True
```
```
>>> out_q5[2] == 'x'
True
```
```
>>> ('carat' in out_q5[3]) or ('x' in out_q5[3])
True
```
```
>>> (out_q5[3] == 'carat * x') or (out_q5[3] == 'x * carat')
True
```
```
>>> s1 = out_q5[4] in ('ordinal_color', 'ordinal_clarity', 'one_hot_clarity_SI2')
>>> s2 = out_q5[4] in ('color', 'clarity')
>>> s1 or s2
True
```
```
>>> (out_q5[4] == 'ordinal_color') or (out_q5[4] == 'color')
True
```
```
>>> np.isclose(out_q5[5], 1434.84, atol=0.1)
True
```


Tests for q6.1 (points: None)
```
>>> isinstance(q6a_out, np.ndarray)
True
```
```
>>> q6a_out[172, 0] == 1
True
```
```
>>> q6a_out[0, 0] == 0
True
```
```
>>> int(q6a_out[172, 0]) == 1
True
```
```
>>> int(q6a_out[0, 0]) == 0
True
```
```
>>> ((q6a_out == 0) | (q6a_out == 1)).all().all()
True
```
```
>>> np.isclose(q6a_out.mean(), 0.32447, atol=0.01)
True
```
```
>>> np.isclose(q6a_out.mean(), 0.32447, atol=0.05)
True
```


Tests for q6.2 (points: None)
```
>>> isinstance(q6b_out, np.ndarray)
True
```
```
>>> np.isclose(q6b_out[0, 0], 0.0075757, atol=0.0001)
True
```
```
>>> np.isclose(q6b_out[1, 0], 0.0025252, atol=0.0001)
True
```
```
>>> np.isclose(np.median(q6b_out), 0.4848, atol=0.02)
True
```
```
>>> np.isclose(q6b_out.squeeze()[-1], 0.5606, atol=0.01)
True
```
```
>>> ((q6b_out >= 0) | (q6b_out <= 1)).all().all()
True
```
```
>>> np.isclose(q6b_out_top_1000.squeeze()[0], 0.01515, atol=0.001)
True
```
```
>>> np.isclose(q6b_out_top_1000.squeeze()[-1], 0.67171, atol=0.01)
True
```
```
>>> ((q6b_out_top_1000 >= 0) | (q6b_out_top_1000 <= 1)).all().all()
True
```


Tests for q6.3 (points: None)
```
>>> len(q6c_out.shape) == 1
True
```
```
>>> np.isclose(q6c_out[0], 61.286, atol=0.0001)
True
```
```
>>> np.isclose(q6c_out[-1], 62.22222, atol=0.001)
True
```
```
>>> np.isclose(np.nanmean(q6c_out), 61.744, atol=10.0)
True
```
```
>>> np.isclose(np.nanmean(q6c_out), 61.744, atol=1.0)
True
```
```
>>> np.isclose(np.nanmax(q6c_out), 619.279, atol=0.1)
True
```
